---
id: 68669dafcb166f535ec05344
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

Now, add an object to the `cardInfo` array.

It should have three keys: `name`, `title`, and `bio`. Assign `"Mark"`, `"Frontend Developer"` and `"I like to work with different frontend technologies and play video games."` to them, respectively.

# --hints--

You should have an object inside of `cardInfo`.

```js
assert.match(__helpers.removeJSComments(code), /const\s*cardInfo\s*=\s*\[\s*\{.*\}\]/s);
```

Your object should have a `name` property with a value of `"Mark"`.

```js
const regex = /const\s*cardInfo\s*=\s*\[\s*\{.*\}\]/s
const matched = __helpers.removeJSComments(code).match(regex);

// Turn the string into an actual object to avoid regex madness

const object = JSON.parse(matched[0].match(/\{.*\}/s));
assert.equal(object.name, "Mark");
```

Your object should have a `title` property with a value of `"Frontend Developer"`.

```js
const regex = /const\s*cardInfo\s*=\s*\[\s*\{.*\}\]/s
const matched = __helpers.removeJSComments(code).match(regex);

const object = JSON.parse(matched[0].match(/\{.*\}/s));
assert.equal(object.title, "Frontend Developer");
```

Your object should have a `bio` property with a value of `"I like to work with different frontend technologies and play video games."`. Make sure to check for any spelling errors.

```js
const regex = /const\s*cardInfo\s*=\s*\[\s*\{.*\}\]/s
const matched = __helpers.removeJSComments(code).match(regex);

const object = JSON.parse(matched[0].match(/\{.*\}/s));
assert.equal(object.bio, "I like to work with different frontend technologies and play video games.");
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Reusable Card component</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>

    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
     <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { App } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <App />
      );
    </script>
  </body>
</html>
```

```css
:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: space-around;
  align-items: center;
}

.card {
  border: 5px solid var(--dark-orange);
  border-radius: 10px;
  width: 100%;
  padding: 20px;
  margin: 10px 0;
  background-color: var(--light-grey);
}

.card-title {
  border-bottom: 4px solid var(--dark-orange);
  width: fit-content;
}

@media (min-width: 768px) {
  .card {
    width: 300px;
  }
}
```

```jsx
export function Card({ name, title, bio }) {
  return (
    <div className="card">
      <h2>{name}</h2>
      <p className="card-title">{title}</p>
      <p>{bio}</p>
    </div>
  )
}

--fcc-editable-region--
export function CardList() {
  const cardInfo = []
  return ""
}
--fcc-editable-region--

export function App() {
  return (
    <div className="flex-container">
      <Card
        name="Mark"
        title="Frontend developer"
        bio="I like to work with different frontend technologies and play video games."
      />

      <Card
        name="Tiffany"
        title="Engineering manager"
        bio="I have worked in tech for 15 years and love to help people grow in this industry."
      />

      <Card
        name="Doug"
        title="Backend developer"
        bio="I have been a software developer for over 20 years and I love working with Go and Rust."
      />
    </div>
  );
}
```
